# Assignment 3
*Topics: difference formulas, finite difference method, interpolation*

In this assignment, we will learn how to use linear algebra to compute numberical solutions of differential equations. The method we will use is called the *finite difference method* and is explained in detail in Section 1.5 of the New Notes (posted on canvas). 

For details on how to upload your notebook back to canvas, check the canvas homepage for a quick tutorial. For additional help with the jupyter workflow and python in general, attend an office hour at 6pm PST on Wednesday.

*Double check your variable names, and don't import extra libraries!*


```python
import numpy as np
from math import *
```

**Exercise 1**

For a differentiable function $f$, recall the *forward*
$$f'(x) \approx \frac{f(x+h)-f(x)}{h},$$
*backward*
$$f'(x) \approx \frac{f(x)-f(x-h)}{h},$$
and *central* 
$$f'(x) \approx \frac{f(x+h)-f(x-h)}{2h},$$
difference formulas. In this exercise we'll compare the above formulas for the function $f(x) = \cos^3(2x)$.

**Exercise 1(a)** Approximate the value of $f'(x)$ at $x = 0,0.2,0.4$ using the forward difference formula and $h = 0.02$. If your approximations are $a,b,c$ for $x = 0,0.2,0.4$ respectively, record you answer in the form FV = [a,b,c]. 

Note: because we have imported all of math, we have the cosine function at our disposal. The code x = cos(0.4) will assign x the value of cos(0.4), which is about 0.92.



<!-- grade_id: cell-273d7c06f6f46f1d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

```

<!-- grade_id: cell-e958f20df6f8bce5, points: 1, grade: true, solution: false -->
```python
assert isinstance(FV,list), "Make sure your FV is a list. Make with square brackets."
print("Checkpoint 1.1 passed: FV data type is good!")


```

**Exercise 1(b)** Approximate the value of $f'(x)$ at $x = 0,0.2,0.4$ using the backward difference formula and $h = 0.02$. If your approximations are $a,b,c$ for $x = 0,0.2,0.4$ respectively, record you answer in the form BV = [a,b,c]. 

<!-- grade_id: cell-f678bc182f382420, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

```

<!-- grade_id: cell-bcae51268cc40df4, points: 1, grade: true, solution: false -->
```python
assert isinstance(FV,list), "Make sure your BV is a list. Make with square brackets."
print("Checkpoint 1.2 passed: BV data type is good!")

```

**Exercise 1(c)** Approximate the value of $f'(x)$ at $x = 0,0.2,0.4$ using the central difference formula and $h = 0.02$. If your approximations are $a,b,c$ for $x = 0,0.2,0.4$ respectively, record you answer in the form CV = [a,b,c]. 

<!-- grade_id: cell-2fc6d2ed4b476f91, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

```

<!-- grade_id: cell-b6f026dbcaff814d, points: 1, grade: true, solution: false -->
```python
assert isinstance(CV,list), "Make sure your BV is a list. Make with square brackets."
print("Checkpoint 1.3 passed: CV data type is good!")

```

**Exercise 1(d)** Calculate the exactly value of $f'(x)$ at $x = 0,0.2,0.4$. Let 

 $$ERF = |FV[0] - f'(0)| + |FV[1] - f'(0.2)| + |FV[2] - f'(0.4)|,$$
 
 be the cumulative error of the forward difference method. Similarly, let 
 
  $$ERB = |BV[0] - f'(0)| + |BV[1] - f'(0.2)| + |BV[2] - f'(0.4)|,$$
  
  and 
  
   $$ERC = |CV[0] - f'(0)| + |CV[1] - f'(0.2)| + |CV[2] - f'(0.4)|,$$
   
   be the cumulative error of the backward and central difference methods. Note which of the methods gave the smallest cumulative error. Assign the variable MERR the value of $\min\{ ERF,ERB,ERC\}$, i.e. the smallest value among $ERF,ERB,ERC$.
   
Note, the python function abs() returns the absolute value of a number. For example abs(-3.1)= 3.1. Also, the python function min() returns the smallest input. For example, min(-10,0.01,2) = -10. 
    
    

<!-- grade_id: cell-2beafa476f2fdeab, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

```

<!-- grade_id: cell-5d773bace2d53ea0, points: 1, grade: true, solution: false -->
```python
assert isinstance(MERR,float), "Make sure your MERR is just a number."
print("Checkpoint 1.4 passed: MERR data type is good!")

```

**Finite Difference Method**

Consider a second order linear differential equation with boundary conditions

$$
y'' + p(t)y' + q(t)y = r(t) \ \ , \ \ y(t_0) = \alpha \ \ , \ \ y(t_f) = \beta
$$

The [finite difference method](https://en.wikipedia.org/wiki/Finite_difference_method) is:

1. Discretize the domain: choose $N$, let $h = (t_f - t_0)/(N+1)$ and define $t_k = t_0 + kh$.
2. Let $y_k \approx y(t_k)$ denote the approximation of the solution at $t_k$.
3. Substitute finite difference formulas into the equation to define an equation at each $t_k$.
4. Rearrange the system of equations into a linear system $A \mathbf{y} = \mathbf{b}$ and solve for
$$
\mathbf{y} = \begin{bmatrix} y_1 & y_2 & \cdots & y_N \end{bmatrix}^T
$$



**General formula** 

*(See "New Notes", p.28 for the following derivation and p.29-30 for an example.)*

Consider the general form of a second order linear ordinary differential equation with boundary conditions
$$
y'' + p(t)y' + q(t)y = r(t) \ \ , \ \ y(t_0) = \alpha \ , \ \ y(t_f) = \beta
$$
Choose $N$ and let  $\displaystyle h = \frac{t_f - t_0}{N+1}$ and define $t_k = t_0 + kh$. Let $y_k$ denote an approximation of $y(t_k)$. Note that the boundary conditions give us $y_0 = \alpha$ and $y_{N+1} = \beta$ and let
$$
\boldsymbol{y} = \begin{bmatrix} y_1 & y_2 & \cdots & y_N \end{bmatrix}^T
$$
Let $p_k = p(t_k)$, $q_k = q(t_k)$ and $r_k = r(t_k)$, and substitute the central difference formulas for both $y''$ and $y'$ at $t_k$ into the differential equation
$$
\frac{y_{k+1} -2y_k + y_{k-1}}{h^2} + p_k \frac{y_{k+1} - y_{k-1}}{2h} + q_k y_k = r_k
$$
Rearrange the equation
\begin{align*}
y_{k+1} -2y_k + y_{k-1} + \frac{h p_k}{2} \left( y_{k+1} - y_{k-1} \right) + h^2q_k y_k &= h^2 r_k \\
\left( 1 - \frac{h p_k}{2} \right) y_{k-1} + (h^2q_k - 2)y_k + \left(1 + \frac{h p_k}{2} \right)y_{k+1} &= h^2 r_k
\end{align*}
Introduce the notation
$$
a_k = 1 - \frac{h p_k}{2}
\hspace{10mm}
b_k = h^2q_k - 2
\hspace{10mm}
c_k = 1 + \frac{h p_k}{2}
$$
Use the boundary conditions $y_0 = \alpha$ and $y_{N+1} = \beta$ and rearrange the equations
$$
\begin{array}{rrrrcrrrrcc}
b_1 y_1 & + & c_1 y_2 & & & & & & & = & h^2 r_1 - \left( 1 - h p_1/2 \right) \alpha \\
a_2 y_1 & + & b_2 y_2 & + & c_2 y_3 & & & & & = & h^2 r_2 \\
& & & & \ddots & & & & & \vdots & \\
& & & & a_{N-1}y_{N-2} & + & b_{N-1}y_{N-1} & + & c_{N-1}y_N & = & h^2 r_{N-1} \\
& & & & & & a_Ny_{N-1} & + & b_Ny_N & = & h^2 r_N - \left( 1 + h p_N/2 \right) \beta
\end{array}
$$
Rewrite in matrix form $A \boldsymbol{y} = \boldsymbol{b}$ where
$$
A =
\left[ \begin{array}{rrcrr}
b_1 & c_1 & & & \\
a_2 & b_2 & c_2 & & \\
& & \ddots & & \\
& & a_{N-1} & b_{N-1} & c_{N-1} \\
& & & a_N & b_N
\end{array} \right]
\hspace{10mm}
\boldsymbol{b} = 
\begin{bmatrix}
h^2 r_1 - \left( 1 - h p_1/2 \right) \alpha \\ h^2 r_2 \\ \vdots \\ h^2 r_{N-1} \\ h^2 r_N - \left( 1 + h p_N/2 \right) \beta
\end{bmatrix}
$$

**Exercise 2** 

In this exercise we'll use the finite difference method to approximate a solution to the second order linear differential equation 

$$ y'' = 4\cos(2x) + 6x, \quad \text{with boundary conditions} \quad y(1) = 2\sin^2(1) + 1, y(3) = 2\sin^2(3) + 27.$$


**Exercise 2(a)** Use the finite difference method with $N = 4$ to find approximate values of $y(1.4),y(1.8),y(2.2),y(2.6)$. 

First you'll need to find the appropriate $4 \times 4$ matrix $A$. Assign the variable $A21$ the np.array matrix $A$. 

Second, record your approximations as the variable A22 = [a,b,c,d] where $a,b,c,d$ are the approximations you found for $y(1.4),y(1.8),y(2.2),y(2.6)$ respectively. 

Both answers can be recorded in the below box, we'll give immediate feedback for each. 



<!-- grade_id: cell-2224af450133b1b7, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

```

<!-- grade_id: cell-9dbf1a9877e57403, points: 2, grade: true, solution: false -->
```python
assert isinstance(A21,np.ndarray), "Check that your A21 is a np.array."
print("Checkpoint 2.1 passed: Good data type.")

assert isinstance(A22,list), "Check that your A22 is a list."
print("Checkpoint 2.2 passed: Good data type.")

```

**Exercise 2(b)** Use the finite difference method with $N = 7$ to find approximate values of $y(1.25),y(1.5),y(1.75),y(2),y(2.25),y(2.5),y(2.75)$. Record your answer as the variable B2 = [a,b,c,d,e,f,g] where $a,b,c,d,e,f,g$ are the approximations you found for $y(1.25),y(1.5),y(1.75),y(2),y(2.25),y(2.5),y(2.75)$ respectively.

<!-- grade_id: cell-0d83951db5bb5fe9, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

```

<!-- grade_id: cell-b7f38aa124445743, points: 2, grade: true, solution: false -->
```python
assert isinstance(B2,list), "Check that your B2 is a list."
print("Checkpoint 2.3 passed: Good data type.")

```

**Exercise 2(c)** Let a,b,c,d,e,f,g be your approximations for $y(1.25),y(1.5),y(1.75),y(2),y(2.25),y(2.5),y(2.75)$ in the above question. Use Vandermonde Interpolation to find a degree 6 polynomial that interpolates the set of points 
$$ \{ (1.25,a), (1.5,b), (1.75,c), (2,d), (2.25,e), (2.5,f), (2.75,g) \}. $$

If the polynomial you find is 
$$ C_0 + C_1 x + C_2 x^2 + C_3 x^3 + C_4 x^4 + C_5 x^5 + C_6 x^6,$$

assign the variable INTPOLY2 the value [C_0,C_1,C_2,C_3,C_4,C_5,C_6]. 

Note that here we are not adding the boundary points of the differential equation into our interpolation.



<!-- grade_id: cell-21a4823159122a7b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

```

<!-- grade_id: cell-bc07a020d9e2f74d, points: 2, grade: true, solution: false -->
```python
assert isinstance(INTPOLY2,list), "Check that your INTPOLY2 is a list."
print("Checkpoint 2.4 passed: Good data type.")

```

**Exercise 2(d)** This differential equation has the exact solution $y(x) = 2\sin^2(x) + x^3$. Lets do a quick calculation to track the average discrepancy between our interpolating polynomial and the exact solution. If $y(x)$ is the exact solution and $p(x)$ is your interpolating polynomial, calculate an average error:

$$ D2 = \frac{1}{4} \left( |p(1.3) - y(1.3)| + |p(1.7) - y(1.7)| + |p(2.3) - y(2.3)| + |p(2.7)-y(2.7)| \right) .$$

Record your answer by assigning the variable E2 the value you calculate for the above expression. 






<!-- grade_id: cell-a8c766f6c311b050, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

```

<!-- grade_id: cell-4de67f6077626ce3, points: 1, grade: true, solution: false -->
```python
assert isinstance(D2,float), "Check that your D2 is a number."
print("Checkpoint 2.5 passed: Good data type.")

```

**Exercise 3** 

This exercise will be similar to Exercise 2. We'll use the finite difference method to approximate a solution to the second order linear differential equation 

$$ y'' + \sin(t) y ' + t^2 y = 4, \quad \text{with boundary conditions} \quad y(0) = 0.5, y'(1) = 3.$$


**Exercise 3(a)** Use the finite difference method with $N = 4$ to find approximate values of $y(0.2),y(0.4),y(0.6),y(0.8),y(1)$. 


First you'll need to find the appropriate $5 \times 5$ matrix $A$. Assign the variable $A31$ the np.array matrix $A$. 

Second, record your approximations as the variable A32 = [a,b,c,d,e] where $a,b,c,d,e$ are the approximations you found for $y(0.2),y(0.4),y(0.6),y(0.8),y(1)$ respectively.

Hint: To use the boundary condition $y'(1) = 4$, use the backward difference formula. We know that
$$ y'(1) \approx \frac{y(1) - y(0.8)}{0.2} ,$$
this is a linear equation that can be added as a row in our `finite difference matrix'.




<!-- grade_id: cell-e8f36e197ebe8c0f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

```

<!-- grade_id: cell-f704bbb8df4d11aa, points: 2, grade: true, solution: false -->
```python
assert isinstance(A31,np.ndarray), "Check that your A31 is a np.array."
print("Checkpoint 3.1 passed: Good data type.")

assert isinstance(A32,list), "Check that your A32 is a list."
print("Checkpoint 3.2 passed: Good data type.")

```

**Exercise 3(b)** Use the finite difference method with $N = 7$ to find approximate values of $y(0.125),y(0.25),y(0.375),y(0.5),y(0.625),y(0.75),y(0.875),y(1)$. Record your answer as the variable B3 = [a,b,c,d,e,f,g,h] where $a,b,c,d,e,f,g,h$ are the approximations you found for $y(0.125),y(0.25),y(0.375),y(0.5),y(0.625),y(0.75),y(0.875),y(1)$ respectively.

<!-- grade_id: cell-7ffefda739e27969, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

```

<!-- grade_id: cell-9eba5774fddcb522, points: 2, grade: true, solution: false -->
```python
assert isinstance(B3,list), "Check that your B2 is a list."
print("Checkpoint 3.3 passed: Good data type.")

```

**Exercise 3(c)** Let a,b,c,d,e,f,g,h be your approximations for $y(0.125),y(0.25),y(0.375),y(0.5),y(0.625),y(0.75),y(0.875),y(1)$ in the above question. Use Vandermonde Interpolation to find a degree 7 polynomial that interpolates the set of points 
$$ \{ (0.125,a), (0.25,b), (0.375,c), (0.5,d), (0.625,e), (0.75,f), (0.875,g), (1,h) \}. $$

If the polynomial you find is 
$$ C_0 + C_1 x + C_2 x^2 + C_3 x^3 + C_4 x^4 + C_5 x^5 + C_6 x^6+C_7x^7,$$

assign the variable INTPOLY3 the value [C_0,C_1,C_2,C_3,C_4,C_5,C_6,C_7]. 

Note here that we are including $x = 1$ in the interpolation, but not $x = 0$. 



<!-- grade_id: cell-bbc22004d0579d8a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

```

<!-- grade_id: cell-b5fec6de53f7217a, points: 2, grade: true, solution: false -->
```python
assert isinstance(INTPOLY3,list), "Check that your INTPOLY2 is a list."
print("Checkpoint 3.4 passed: Good data type.")

```